# MATH 307 Assignment 4

* See [Mathematical Python](https://www.math.ubc.ca/~pwalls/math-python/) for an introduction to Python and Jupyter
* See [MATH 307 Jupyter Notebooks](https://ubcmath.github.io/MATH307/) for examples
* Run the tests to verify your solutions but note there are **hidden tests**
* Submit your `.ipynb` notebook file to Canvas

```python
import numpy as np
import scipy.linalg as la
import scipy.signal as ss
import matplotlib.pyplot as plt
from scipy.io import wavfile
from IPython.display import Audio
plt.rcParams["figure.figsize"] = (18,5)
```

## Problem 1 (6 marks)

The code below creates a square wave with frequency $k$ periods per second using $N$ samples per second.

```python
N = 100
t = np.arange(N)/N
k = 2
x_square = ss.square(2*np.pi*k*t)
plt.stem(x_square)
plt.grid(True)
plt.show()
```

(a) Create a square wave with frequency $k=8$ periods per second using with $N=128$ samples per second. Save the result as `x1`.

<!-- grade_id: cell-fe0d446dc88269fe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x1 = 
```

<!-- grade_id: cell-1ce313ce2e1c55f8, points: 1, grade: true, solution: false -->
```python
"Verify x1 is a NumPy array. (1 mark)"
assert isinstance(x1,np.ndarray)
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-270a334d5d47637e, points: 1, grade: true, solution: false -->
```python
"Verify the values of x1. This cell contains hidden tests. (1 mark)"
assert np.allclose(x1[6:10],[1,1,-1,-1])
print("Problem 1a Test 2: Success!")
```

Visualize the signal:

```python
plt.stem(x1)
plt.grid(True)
plt.show()
```

(b) Compute the discrete Fourier transform of `x1` and save the result as `y1`.

<!-- grade_id: cell-80f837d7588ff005, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y1 = 
```

<!-- grade_id: cell-22538064b1f82e97, points: 1, grade: true, solution: false -->
```python
"Verify y1 is a NumPy array with complex values. (1 mark)"
assert isinstance(y1,np.ndarray) and isinstance(y1[0],np.complex128)
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-2d004d0e87ed9f31, points: 1, grade: true, solution: false -->
```python
"Verify the values of y1. This cell contains hidden tests. (1 mark)"
assert np.allclose(y1[1],1.0823922+2.61312593j)
print("Problem 1a Test 2: Success!")
```

Visualize the result:

```python
plt.stem(np.abs(y1))
plt.grid(True)
plt.show()
```

(c) The square wave `x1` is dominated by 4 main frequencies. Save the list of integers in increasing order (like `[1,2,3,4,5]`) as `x1_freqs`.

<!-- grade_id: cell-9deceba60b49a7a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x1_freqs = 
```

<!-- grade_id: cell-931ec78e470dd627, points: 1, grade: true, solution: false -->
```python
"Verify x1_freqs is a list. (1 mark)"
assert isinstance(x1_freqs,list)
print("Problem 1c Test 1: Success!")
```

<!-- grade_id: cell-47ffd7ea88941b6f, points: 1, grade: true, solution: false -->
```python
"Verify the values of x1_freqs. This cell contains hidden tests. (1 mark)"
assert np.allclose(x1_freqs[0],8)
print("Problem 1c Test 2: Success!")
```

Just for fun ... the code below creates a C major 7 chord with square waves!

```python
N = 44100
t = np.arange(N)/N
ks = [261,330,392,494]
chord = sum([ss.square(2*np.pi*k*t) for k in ks])
Audio(chord,rate=N)
```

## Problem 2 (5 marks)

The cell below loads a sound file `example.wav` as a NumPy array `x2`, records the sample rate `rate` and then allows us to play the sound in the notebook. Note that the sound has some noise and our goal is to filter it out.

```python
rate,x2 = wavfile.read('example.wav')

Audio(x2,rate=rate)
```

(a) Compute the discrete Fourier transform of `x2` and save the result as `y2`.

<!-- grade_id: cell-21050c278f0086c3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y2 = 
```

<!-- grade_id: cell-761538d25b11ed68, points: 1, grade: true, solution: false -->
```python
"Verify y2 is a NumPy array with complex values. (1 mark)"
assert isinstance(y2,np.ndarray) and isinstance(y2[0],np.complex128)
print("Problem 2a Test 1: Success!")
```

<!-- grade_id: cell-84fd5eca0a3d79bc, points: 2, grade: true, solution: false -->
```python
"Verify the values of y2. This cell contains hidden tests. (2 marks)"
assert np.allclose(y2[0],-865729.2483320391)
print("Problem 2a Test 2: Success!")
```

(b) The cells below plot the magnitude of the first 300 frequencies in the Fourer transform `y2` and `y2_filtered`. Compute the inverse Fourier transform of `y2_filtered` and use `np.real` to convert the entries to real numbers and save the result as `x2_filtered`. Play the audio and see if it sounds less noisy.

```python
plt.stem(np.abs(y2[:300]))
plt.grid(True)
plt.show()
```

```python
y2_filtered = y2.copy()
y2_filtered[np.abs(y2) < 1e7] = 0.0

plt.stem(np.abs(y2_filtered[:300]))
plt.grid(True)

plt.show()
```

<!-- grade_id: cell-30a62d6e001779d1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x2_filtered = 
```

<!-- grade_id: cell-a0b0c2e750f77a0a, points: 1, grade: true, solution: false -->
```python
"Verify x2_filtered is a NumPy array with real values. (1 mark)"
assert isinstance(x2_filtered,np.ndarray) and isinstance(x2_filtered[0],np.float64)
print("Problem 2b Test 1: Success!")
```

<!-- grade_id: cell-27db3a40aeeefc4e, points: 1, grade: true, solution: false -->
```python
"Verify the values of x2_filtered. This cell contains hidden tests. (1 mark)"
assert np.allclose(x2_filtered[0],-2459.67554688)
print("Problem 2b Test 2: Success!")
```

```python
Audio(x2_filtered.real,rate=rate)
```